# Testes sintéticos para promtool (regras acima)
# Rodar: promtool test rules ops/watchers/tests/a110_rules_test.yaml

tests:
- interval: 1m
  # Simula buckets de latência para swap crescendo ao longo do tempo
  input_series:
  - series: 'amm_op_latency_seconds_bucket{le="0.06",op="swap",env="dev",service="trendmarket-amm",version="v0.1.0"}'
    values: '0+0x5 1+1x20'
  - series: 'amm_op_latency_seconds_bucket{le="+Inf",op="swap",env="dev",service="trendmarket-amm",version="v0.1.0"}'
    values: '0+0x5 3+3x20'
  alert_rule_test:
  - eval_time: 15m
    alertname: OBS_P95_Swap_TooHigh
    exp_alerts:
    - exp_labels: { severity: 'warn', team: 'ce' }


- interval: 1m
  # Freshness >5s sustentada para oracle
  input_series:
  - series: 'data_freshness_seconds{source="oracle",domain="pricing",service="trendmarket-amm",env="dev"}'
    values: '2x5 6x10'
  alert_rule_test:
  - eval_time: 10m
    alertname: OBS_Oracle_Stale
    exp_alerts:
    - exp_labels: { severity: 'warn', team: 'ce', source: 'oracle' }


- interval: 1m
  # CDC lag >10s por 10m
  input_series:
  - series: 'cdc_lag_seconds{stream="orders",partition="0",service="trendmarket-amm",env="dev"}'
    values: '4x5 12x15'
  alert_rule_test:
  - eval_time: 15m
    alertname: OBS_CDC_Lag_Warn
    exp_alerts:
    - exp_labels: { severity: 'warn', team: 'ce', stream: 'orders' }


- interval: 1m
  # CDC lag crítico >30s por 5m
  input_series:
  - series: 'cdc_lag_seconds{stream="orders",partition="1",service="trendmarket-amm",env="dev"}'
    values: '5x5 35x10'
  alert_rule_test:
  - eval_time: 10m
    alertname: OBS_CDC_Lag_Crit
    exp_alerts:
    - exp_labels: { severity: 'critical', team: 'ce', stream: 'orders' }


- interval: 1m
  # Drift warn >0.2 por 30m
  input_series:
  - series: 'drift_score{feature="price_psi",service="trendmarket-amm",env="dev"}'
    values: '0.1x20 0.25x40'
  alert_rule_test:
  - eval_time: 50m
    alertname: OBS_Drift_Critical_Feature_Warn
    exp_alerts:
    - exp_labels: { severity: 'warn', team: 'ce', feature: 'price_psi' }


- interval: 1m
  # Drift crítico >0.4 por 15m
  input_series:
  - series: 'drift_score{feature="volume_psi",service="trendmarket-amm",env="dev"}'
    values: '0.3x10 0.5x20'
  alert_rule_test:
  - eval_time: 25m
    alertname: OBS_Drift_Critical_Feature_Crit
    exp_alerts:
    - exp_labels: { severity: 'critical', team: 'ce', feature: 'volume_psi' }


- interval: 1m
  # Cobertura de hooks <0.95 por 30m
  input_series:
  - series: 'hook_coverage_ratio{env="dev"}'
    values: '1x10 0.9x40'
  alert_rule_test:
  - eval_time: 40m
    alertname: OBS_Hook_Coverage_Low
    exp_alerts:
    - exp_labels: { severity: 'warn', team: 'ce', env: 'dev' }


- interval: 1m
  # Hook sem execuções em 1h
  input_series:
  - series: 'hook_executions_total{hook_id="hook_pre_trade",status="success",env="dev"}'
    values: '0+0x61'
  alert_rule_test:
  - eval_time: 1h
    alertname: OBS_Hook_Execution_Stalled
    exp_alerts:
    - exp_labels: { severity: 'warn', team: 'ce', hook_id: 'hook_pre_trade' }


- interval: 1m
  # Cardinalidade excedendo 70% do budget
  input_series:
  - series: 'observability_series_budget_ratio{metric="amm_op_latency_seconds_bucket"}'
    values: '0.6x5 0.8x20'
  alert_rule_test:
  - eval_time: 20m
    alertname: OBS_Cardinality_Budget_Warn
    exp_alerts:
    - exp_labels: { severity: 'warn', team: 'ce', metric: 'amm_op_latency_seconds_bucket' }


- interval: 1m
  # Cardinalidade ultrapassando 90% do budget
  input_series:
  - series: 'observability_series_budget_ratio{metric="data_freshness_seconds"}'
    values: '0.5x5 0.95x20'
  alert_rule_test:
  - eval_time: 20m
    alertname: OBS_Cardinality_Budget_Crit
    exp_alerts:
    - exp_labels: { severity: 'critical', team: 'ce', metric: 'data_freshness_seconds' }


- interval: 1m
  # Synthetic ok ratio abaixo de 99%
  input_series:
  - series: 'synthetic_ok_ratio{route="swap",service="trendmarket-amm",env="dev"}'
    values: '1x5 0.95x20'
  alert_rule_test:
  - eval_time: 20m
    alertname: OBS_Synthetic_OK_Ratio_Low
    exp_alerts:
    - exp_labels: { severity: 'warn', team: 'ce', route: 'swap', service: 'trendmarket-amm', env: 'dev' }
